## Exercise 02.1 (if-else)

Consider the following assessment criteria which map a score out of 100 to an 
assessment grade:

| Grade            | Raw score  (/100)      |
| ---------------- | ---------------------- |
| Excellent        | $\ge 85$               |
| Very good        | $\ge 76.5$ and $< 85$  |
| Good             | $\ge 64$ and $< 76.5$  |
| Need improvement | $\ge 40$ and $< 64$    |
| Did you try?     | $< 40$                 |

Write a program that, given an a score, prints the appropriate grade. Print an error message if the input score is greater than 100 or less than zero.

<!-- grade_id: cell-71121f470be22e4a, points: 0, grade: false, solution: true -->
```python
# Score from user
score = 72

if score < 0 or score > 100:
    print("Error: Score must be between 0 and 100.")
else:
    # getting grade based on the score
    if score >= 85:
        print("Grade: Excellent")
    elif score >= 76.5:
        print("Grade: Very good")
    elif score >= 64:
        print("Grade: Good")
    elif score >= 40:
        print("Grade: Need improvement")
    else:
        print("Grade: Did you try?")
```

## Exercise 02.2 (bisection)

Bisection is an iterative method for finding approximate roots of a function. Say we know that the function $f(x)$ has one root between $x_{0}$ and $x_{1}$ ($x_{0} < x_{1}$). We then:

- Evaluate $f$ at the midpoint $x_{\rm mid} = (x_0 + x_1)/2$, i.e. compute
   $f_{\rm mid} = f(x_{\rm mid})$
- Evaluate $f(x_0) \cdot f(x_{\rm mid})$

  - if $f(x_0) \cdot f(x_{\rm mid}) < 0$: 

    $f$ must change sign somewhere between $x_0$ and $x_{\rm mid}$, hence the root must lie between 
    $x_0$ and $x_{\rm mid}$, so set $x_1 = x_{\rm mid}$.
   
  - else:

    $f$ must change sign somewhere between $x_{\rm mid}$ and $x_1$, so set
    $x_0 = x_{\rm mid}$.

The above steps can be repeated a specified number of times, or until $|f_{\rm mid}|$
is below a tolerance, with $x_{\rm mid}$ being the approximate root.


### Task

The function

$$
f(x) = \frac{x^{5}}{10} + x^3 - 10x^2 + 4x + 7
$$


has one root in the range $0 < x < 2$.

1. Use the bisection method to find an approximate root $x_{r}$ using 20 iterations 
   (use a `for` loop).
2. Use the bisection method to find an approximate root $x_{r}$ such that 
   $\left| f(x_{r}) \right| < 1 \times 10^{-6}$ and report the number of iterations 
   required (use a `while` loop).

Store the approximate root using the variable `x_mid`, and store $f(x_{\rm mid})$ using the variable `f`.

*Hint:* Use  `abs` to compute the absolute value of a number, e.g. `y = abs(x)` assigns the absolute value of `x` to `y`. 

#### (1) Using a `for` loop.

<!-- grade_id: cell-3a0d0a59613d6b3f, points: 0, grade: false, solution: true -->
```python
# Initial end points
x0 = 0.0
x1 = 2.0

# Use 20 iterations
for n in range(20):
    # Compute midpoint
    x_mid = (x0 + x1) / 2

    # Evaluate function at (i) left end-point and at (ii) midpoint
    f0 = (x0**5) / 10 + x0**3 - 10 * x0**2 + 4 * x0 + 7
    f = (x_mid**5) / 10 + x_mid**3 - 10 * x_mid**2 + 4 * x_mid + 7
    
     # checking sign
    if f0 * f < 0:
        x1 = x_mid
    else:
        x0 = x_mid

    print(n, x_mid, f)
```

Output:
```
0 1.0 2.0999999999999996
1 1.5 -5.365625000000001
2 1.25 -1.36669921875
3 1.125 0.4477813720703132
4 1.1875 -0.4408627510070797
5 1.15625 0.008327355980872753
6 1.171875 -0.21507753478363156
7 1.1640625 -0.1030741602036862
8 1.16015625 -0.04729774910292761
9 1.158203125 -0.019466230897839054
10 1.1572265625 -0.005564689502373099
11 1.15673828125 0.0013825210451265946
12 1.156982421875 -0.0020907873792292975
13 1.1568603515625 -0.00035405894194262544
14 1.15679931640625 0.0005142496094645566
15 1.156829833984375 8.009997302949046e-05
16 1.1568450927734375 -0.00013697832466341708
17 1.1568374633789062 -2.8438885864900953e-05
18 1.1568336486816406 2.5830616070976475e-05
19 1.1568355560302734 -1.304116775457942e-06
```

<!-- grade_id: cell-421824bbd4e61fae, points: 0, grade: true, solution: false -->
```python
## tests ##
import math
assert math.isclose(x_mid, 1.1568355560302734)
assert abs(f) < 1e-5
```

#### (2) Using a `while` loop

Use the variable `counter` for the iteration number. 

*Remember to guard against infinite loops.*

<!-- grade_id: cell-1747a3fb1dd1de48, points: 0, grade: false, solution: true -->
```python
# Initial end points
x0 = 0.0
x1 = 2.0

tol = 1.0e-6
error = tol + 1.0

# Iterate until tolerance is met
counter = 0
while error > tol:
    
    x_mid = (x0 + x1) / 2

    # Evaluate function at (i) left end-point and at (ii) midpoint
    f0 = (x0**5) / 10 + x0**3 - 10 * x0**2 + 4 * x0 + 7
    f = (x_mid**5) / 10 + x_mid**3 - 10 * x_mid**2 + 4 * x_mid + 7

    # checking sign
    if f0 * f < 0:
        x1 = x_mid
    else:
        x0 = x_mid

    # get error (absolute value of f at midpoint)
    error = abs(f)
    counter += 1

    # Guard against an infinite loop
    if counter > 1000:
        print("Oops, iteration count is very large. Breaking out of while loop.")
        break
    
    print(counter, x_mid, error)
```

Output:
```
1 1.0 2.0999999999999996
2 1.5 5.365625000000001
3 1.25 1.36669921875
4 1.125 0.4477813720703132
5 1.1875 0.4408627510070797
6 1.15625 0.008327355980872753
7 1.171875 0.21507753478363156
8 1.1640625 0.1030741602036862
9 1.16015625 0.04729774910292761
10 1.158203125 0.019466230897839054
11 1.1572265625 0.005564689502373099
12 1.15673828125 0.0013825210451265946
13 1.156982421875 0.0020907873792292975
14 1.1568603515625 0.00035405894194262544
15 1.15679931640625 0.0005142496094645566
16 1.156829833984375 8.009997302949046e-05
17 1.1568450927734375 0.00013697832466341708
18 1.1568374633789062 2.8438885864900953e-05
19 1.1568336486816406 2.5830616070976475e-05
20 1.1568355560302734 1.304116775457942e-06
21 1.156834602355957 1.2263254177469207e-05
22 1.1568350791931152 5.479569834321296e-06
23 1.1568353176116943 2.087726812760593e-06
24 1.1568354368209839 3.9180508970559913e-07
```

<!-- grade_id: cell-ddff32f44df88388, points: 0, grade: true, solution: false -->
```python
## tests ##
assert counter == 24
assert abs(f) < 1.0e-6
```

## Exercise 02.3 (series expansion)

For $|x| < 1$ the series: 

$$
(1 + x)^{-1/2} = \sum_{n = 0}^{\infty} \frac{(-1)^n (2n)!}{4^n (n!)^2} x^n
$$

converges.

1. Using a `for` statement, approximate $1/\sqrt{0.16}$ using 30 terms in the series expansion and report the absolute error.

1. Using a `while` statement, compute how many terms in the series are required to approximate $1/\sqrt{0.16}$ to within $1 \times 10^{-5}$. 

Store the absolute value of the error in the variable `error`.

### Hints

To compute the factorial, use the Python `math` module:
```python
import math
nfact = math.factorial(10)
```

You only need `import math` once at the top of your program. Standard modules, like `math`, will be explained in a later

<!-- The power series expansion for the sine function is: 

$$
\sin(x) = \sum_{n = 0}^{\infty} (-1)^n \frac{x^{2n +1}}{(2n+1)!}
$$

(See mathematics data book for a less compact version; this compact version is preferred here as it is simpler to program.)

1. Using a `for` statement, approximate $\sin(3\pi/2)$ using 15 terms in the series expansion and report the absolute error.

1. Using a `while` statement, compute how many terms in the series are required to approximate $\sin(3\pi/2)$ to within $1 \times 10^{-8}$. 

Store the absolute value of the error in the variable `error`.

*Note:* Calculators and computers use iterative or series expansions to compute trigonometric functions, similar to the one above (although they use more efficient formulations than the above series).

### Hints

To compute the factorial and to get a good approximation of $\pi$, use the Python `math` module:
```python
import math
nfact = math.factorial(10)
pi = math.pi
```
You only need '`import math`' once at the top of your program. Standard modules, like `math`, will be explained in a later. If you want to test for angles for which sine is not simple, you can use 
```python
a = 1.3
s = math.sin(a)
```    
to get an accurate computation of sine to check the error. -->

#### (1) Using a `for` loop

<!-- grade_id: cell-d6df4b428f969a5d, points: 0, grade: false, solution: true -->
```python
# Import the math module to access math.factorial
import math

# Value of x (such that (1 + x) = 0.16  
x = -0.84

# Initialise approximation of the function
approx_f = 0.0

# Adding first 30 terms of series
for n in range(30):
    term = ((-1)**n) * math.factorial(2*n) / (4**n * (math.factorial(n)**2)) * (x**n)
    approx_f += term

true_value = 1 / math.sqrt(0.16)

error = abs(approx_f - true_value)
    
print("The error is:")
print(error)
```

Output:
```
The error is:
0.0031921895736983785
```

<!-- grade_id: cell-8b269b717e92414d, points: 0, grade: true, solution: false -->
```python
## test ##
assert error >= 0
assert error < 1.0e-2
```

#### (2) Using a `while` loop

*Remember to guard against infinite loops.*

<!-- grade_id: cell-03df25d83df3674e, points: 0, grade: false, solution: true -->
```python
# Import the math module to access math.sin and math.factorial
import math

# Value of x (such that (1 - x) = 0.16)
x = -0.84

# Tolerance and initial error (this just needs to be larger than tol)
tol = 1.0e-5
error = tol + 1.0

# Initialise approximation of function
approx_f = 0.0
true_value = 1 / math.sqrt(0.16)
# Initialise counter
n = 0

# Loop until error satisfies tolerance, with a check to avoid 
# an infinite loop
while error > tol and n < 1000:
        
    term = ((-1)**n) * math.factorial(2*n) / (4**n * (math.factorial(n)**2)) * (x**n)
    approx_f += term
    error = abs(approx_f - true_value)
    # Increment counter
    n += 1    
    
    
print("\nThe error is:", error)
print("Number of terms in series:", n)
```

Output:
```
The error is: 8.689878570500298e-06
Number of terms in series: 62
```

<!-- grade_id: cell-03df25d83df3674dfwe, points: 0, grade: true, solution: false -->
```python
## test ##
assert error >= 0
assert error <= 1.0e-5
```