<!-- BEGIN QUESTION -->

## Exercise 03.1

Compare the computed values of 

$$ 
d_0 = a \cdot b + a  \cdot  c
$$

and 

$$ 
d_1 = a \cdot (b + c)
$$

when $a = 100$, $b = 0.1$ and $c = 0.2$. Store $d_{0}$ in the variable `d0` and $d_{1}$ in the variable `d1`.


Try checking for equality, e.g. `print(d0 == d1)`. 

<!-- grade_id: cell-447110c225124fc4, points: 0, grade: false, solution: true -->
```python
a = 100
b = 0.1
c = 0.2

d0 = a*b + a*c
print(d0)

d1 = a*(b+c)
print(d1)

print(d0 == d1)
```

Output:
```
30.0
30.000000000000004
False
```

<!-- grade_id: cell-91f2d89d88ea86c0, points: 0, grade: true, solution: false -->
```python
## tests ##
assert d0 == 30.0
assert d1 != 30.0
assert d0 != d1
```

<!-- END QUESTION -->

## Exercise 03.2

For the polynomial 
\begin{align}
f(x, y) &= (x + y)^{6} 
\\
&=  x^6 + 6x^{5}y + 15x^{4}y^{2} + 20x^{3}y^{3} + 15x^{2}y^{4} + 6xy^{5} + y^{6}
\end{align}
compute $f$ using: (i) the compact form $(x + y)^{6}$; and (ii) the expanded form for:

(a) $x = 10$ and $y = 10.1$

(b) $x = 10$ and $y = -10.1$

and compare the number of significant digits for which the answers are the same.
Store the answer for the compact version using the variable `f0`, and using the variable `f1` for the expanded version.

For case (b), compare the computed and analytical solutions and consider the relative error.
Which approach would you recommend for computing this expression?

#### (a) $x = 10$ and $y = 10.1$

<!-- grade_id: cell-e8428e3de87d5500, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = 10.1

f0 = (x + y) ** 6

f1 = x ** 6 + 6 * (x ** 5) * y + 15 * (x ** 4) * (y**2) + 20 * (x ** 3) * (y ** 3) + 15 * (x ** 2) * (y ** 4) + 6 * x * (y ** 5) + (y ** 6)


```

<!-- grade_id: cell-55d184514edf3a67, points: 0, grade: true, solution: false -->
```python
## tests ##
import math
assert math.isclose(f0, 65944160.60120103, rel_tol=1e-10)
assert math.isclose(f1, 65944160.601201, rel_tol=1e-10)
```

#### (b) $x = 10$ and $y = -10.1$

<!-- grade_id: cell-1561d20281ff67f6, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = -10.1

f0 = (x + y) ** 6

f1 = x ** 6 + 6 * (x ** 5) * y + 15 * (x ** 4) * (y**2) + 20 * (x ** 3) * (y ** 3) + 15 * (x ** 2) * (y ** 4) + 6 * x * (y ** 5) + (y ** 6)

#here we need to use compact form since near cancellation of 10 and 10.1 loses precision hence insignificant digits cause magnified error. 
```

```python
## tests ##
import math
assert math.isclose(f0, 1.0e-6, rel_tol=1e-10)
assert math.isclose(f1, 1.0e-6, rel_tol=1e-2)
```

## Exercise 03.3

Consider the expression

$$
f = \frac{1}{\sqrt{x^2 - 1} - x}
$$

When $x$ is very large, the denominator approaches zero, which can cause problems.

Try rephrasing the problem and eliminating the fraction by multiplying the numerator and denominator by $\sqrt{x^2 - 1} + x$ and evaluate the two versions of the expression when:

(a) $x = 1 \times 10^{7}$

(b) $x = 1 \times 10^{9}$ (You may get a Python error for this case. Why?)

#### (a) $x = 1 \times 10^{7}$

<!-- grade_id: cell-149ac8791e7ff020, points: 0, grade: true, solution: true -->
```python
import math
x = 1.0e7

f_orig = 1.0 / (math.sqrt(x*x - 1.0) - x)

f_modified = -(math.sqrt(x*x - 1.0) + x)

print(f_orig, f_modified)
```

Output:
```
-19884107.85185185 -19999999.999999948
```

#### (b) $x = 1 \times 10^{9}$

<!-- grade_id: cell-4ef986ef1f5838ba, points: 0, grade: true, solution: true -->
```python
import math
x = 1.0e9

f_orig = 1.0 / (math.sqrt(x*x - 1.0) - x)

f_modified = -(math.sqrt(x*x - 1.0) + x)

print(f_orig, f_modified)

#Error occurs due to x being large enough to give ZeroDivisionError, since it not possible to divide by zero.
```

Output:
```
ZeroDivisionError                         Traceback (most recent call last)
Cell In[29], line 4
      1 import math
      2 x = 1.0e9
----> 4 f_orig = 1.0 / (math.sqrt(x*x - 1.0) - x)
      6 f_modified = -(math.sqrt(x*x - 1.0) + x)
      8 print(f_orig, f_modified)

ZeroDivisionError: float division by zero
```