## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc, to develop a new version of the function with an improved interface (the interface is the arguments that are passed to the function) with fewer arguments. Add appropriate checks inside your function on the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
def area_triangle(vertices):
    if not (isinstance(vertices, list) and len(vertices) == 3):
        raise ValueError("Vertices must be a list of three co-ordinates")

    for v in vertices:
        if not (isinstance(v, tuple) and len(v) == 2):
            raise ValueError("Each vertex must be a tuple of length 2")
    x0, y0 = vertices[0]
    x1, y1 = vertices[1]
    x2, y2 = vertices[2]
    
    return abs((x0*(y1 - y2) + x1*(y2 - y0) + x2*(y0 - y1)) / 2)
```

```python
# testing

test = [(0,0), (0,2), (3,0)]
assert area_triangle(test) == 3.0 
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that, given only the vertex coordinates, computes the area of a simple polygon. Test your function for some simple shapes. Pay careful attention to the range of any loops. **Test for polygons without a vertex at the origin.**

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
def area_polygon(vertices):
    n = len(vertices)
    if n < 3:
        raise ValueError("Polygon must have at least 3 vertices")

    A = 0
    for i in range(n):
        x_i, y_i = vertices[i]
        if i == n-1:
            i_next = 0
        else:
            i_next = i + 1
        x_next, y_next = vertices[i_next]
        A += (x_i * y_next - x_next * y_i)

    return abs(A) / 2
```

```python
# testing

square = [(1,1), (1,3), (4,3), (4,1)]
assert area_polygon(square) == 6.0    #since 3x2 rectangle
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match. Print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    """Return sum of two vectors"""
    if len(x) != len(y):
        raise ValueError(f"Length mismatch, length of x is {len(x)} and length of y is {len(y)}")
    z = [0]*len(x)
    for i in range(len(x)):
        z[i] = x[i] + y[i]

    return z
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
# tests ##
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Optional (advanced) 

Try writing a one-line version that performs the `sum_vector` operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
...
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A Python dictionary from college abbreviation to name; and
1. A Python list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
colleges = {
    "St Catharine's": "CTH",
    "Trinity College": "T",
    "King's College": "K",
    "St John's College": "JN",
    "Christ's College": "CHR",
    "Queens' College": "Q"
}

# abbreviation -> name
abbr_to_name = {}
for key, value in colleges.items():
    abbr_to_name[value] = key

print(abbr_to_name)

sorted_list = []
for key, value in colleges.items():
    sorted_list.append(value)
sorted_list.sort()


print(sorted_list)
    
```

Output:
```
{'CTH': "St Catharine's", 'T': 'Trinity College', 'K': "King's College", 'JN': "St John's College", 'CHR': "Christ's College", 'Q': "Queens' College"}
['CHR', 'CTH', 'JN', 'K', 'Q', 'T']
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
...
```